<!DOCTYPE html>
<html>
<head>
<title>HTML Templates: Clearing stack back to a table context</title>
<meta name="author" title="Sergey G. Grekhov" href="mailto:sgrekhov@unipro.ru">
<meta name="reviewer" title="" href="mailto:"> <!-- 2013-06-20 -->
<meta name="assert" content="Clearing the stack back to a table context must be aborted if the current node is template">
<link rel="help" href="http://www.w3.org/TR/2013/WD-html-templates-20130214/#clearing-the-stack">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src='../../testcommon.js'></script>
<link rel="stylesheet" href="/resources/testharness.css">
</head>
<body>
<div id="log"></div>
<script type="text/javascript">

function doTest(d, templateInnerHTML, id, tagName) {
    
    d.body.innerHTML = '' +
        '<table id="tbl">' +
        '<template id="tmpl1">' +
        // When parser meets <caption>, <colgroup>, <tbody>, <tfoot>, <thead>, <col>
        // stack must be cleared back to table context. But <template> tag should abort this        
        templateInnerHTML + 
        '</template>' +
        '<tr id="tr">' +
        '<td id="td">' +
        '</td>' +
        '</tr>' + 
        '</table>';
    
    var tbl = d.querySelector('#tbl');
    var tr = d.querySelector('#tr');
    var td = d.querySelector('#td');
    var t = d.querySelector('#tmpl1');
    
    assert_equals(tbl.rows.length, 1, 'Wrong number of table rows');
    assert_equals(tbl.rows[0].cells.length, 1, 'Wrong number of table cells');
    assert_equals(t.parentNode, tbl, 'Wrong template parent');
    assert_true(t.content.querySelector('#' + id) != null, 
                'Element should present in the template content');
    assert_equals(t.content.querySelector('#' + id).tagName, tagName, 
                'Wrong element in the template content');
}

//test caption
test(function () {
    var d = newHTMLDocument();
    
    doTest(d, '<caption id="cptn">Table caption</caption>', 'cptn', 'CAPTION');
    
    assert_true(d.querySelector('#tbl').caption == null, 'Table should have no caption');    
    
}, 'Clearing stack back to a table context: <caption>');

//test colgroup
test(function () {
    var d = newHTMLDocument();
    
    doTest(d, '<colgroup id="cgp" width="100%"/>', 'cgp', 'COLGROUP');
        
}, 'Clearing stack back to a table context: <colgroup>');


//test tbody
test(function () {
    var d = newHTMLDocument();
    
    doTest(d, '<tbody id="tbd"></tbody>', 'tbd', 'TBODY');
    
    assert_equals(d.querySelector('#tbl').tBodies.length, 1, 'Table should have no more than one body');
    
}, 'Clearing stack back to a table context: <tbody>');


//test tfoot
test(function () {
    var d = newHTMLDocument();
    
    doTest(d, '<tfoot id="tft"></tfoot>', 'tft', 'TFOOT');
        
    assert_true(d.querySelector('#tbl').tFoot == null, 'Table should have no footer');
    
}, 'Clearing stack back to a table context: <tfoot>');


//test thead
test(function () {
    var d = newHTMLDocument();
    
    doTest(d, '<thead id="thd"></thead>', 'thd', 'THEAD');
    
    assert_true(d.querySelector('#tbl').tHead == null, 'Table should have no header');
    
}, 'Clearing stack back to a table context: <thead>');



//test col
test(function () {
    var d = newHTMLDocument();
    
    doTest(d, '<col id="cl" width="100%"/>', 'cl', 'COL');
        
}, 'Clearing stack back to a table context: <col>');

</script>
</body>
</html>
