<!DOCTYPE html>
<html>
<head>
<title>HTML Templates: The template contents owner document type is HTML document</title>
<meta name="author" title="Sergey G. Grekhov" href="mailto:sgrekhov@unipro.ru">
<meta name="reviewer" title="" href="mailto:"> <!-- 2013-06-20 -->
<meta name="assert" content="The template contents owner document type is HTML document, if template is declared in HTML document">
<link rel="help" href="http://www.w3.org/TR/2013/WD-html-templates-20130214/#definitions">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src='../testcommon.js'></script>
<link rel="stylesheet" href="/resources/testharness.css">
</head>
<body>
<div id="log"></div>
<script type="text/javascript">

// document has browsing context
testInIFrame(null, function(ctx) {
    var d = ctx.iframes[0].contentDocument;
    var t = d.createElement('template');

    var div = d.createElement('div');
    div.setAttribute('id', 'div1');
    div.innerText = 'Some text';

    t.appendChild(div);

    d.body.appendChild(t);

    assert_equals(t.content.ownerDocument.constructor.name, 'HTMLDocument',
            'Template content owner should be a HTML document');

}, 'The template contents owner document type (document has browsing context)');


// document has no browsing context
test(function() {
    var d = newHTMLDocument();
    var t = d.createElement('template');

    d.body.appendChild(t);

    assert_equals(t.content.ownerDocument.constructor.name, 'HTMLDocument',
            'Template content owner should be a HTML document');

}, 'The template contents owner document type (document has no browsing context)');
</script>
</body>
</html>
