<!DOCTYPE html>
<html>
<head>
<title>HTML Templates: The template contents owner document (no browsing context)</title>
<meta name="author" title="Sergey G. Grekhov" href="mailto:sgrekhov@unipro.ru">
<meta name="reviewer" title="Aleksei Yu. Semenov" href="mailto:a.semenov@unipro.ru"> <!-- 2013-06-26 -->
<meta name="assert" content="If template's enclosing document has no browsing context, it must be template contents owner">
<link rel="help" href="http://www.w3.org/TR/2013/WD-html-templates-20130214/#definitions">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src='../testcommon.js'></script>
<link rel="stylesheet" href="/resources/testharness.css">
</head>
<body>
<div id="log"></div>
<script type="text/javascript">

//document content owner is the current document (has no browsing content)
test(function() {
    var d = newHTMLDocument();
    var t = d.createElement('template');

    d.body.appendChild(t);

    assert_equals(t.content.ownerDocument, d, 'Wrong template content owner');

}, 'The template contents owner document (no browsing context)');

//document content owner is the current document (has no browsing content)
// using innerHTML to create template
test(function() {
    var d = newHTMLDocument();

    d.body.innerHTML = '<template><div>some text</div></template>';

    var t = d.querySelector('template');

    assert_equals(t.content.ownerDocument, d, 'Wrong template content owner');

}, 'The template contents owner document (no browsing context) using innerHTML');

</script>
</body>
</html>
