<!DOCTYPE html>
<html>
<head>
<title>HTML Templates: Template element's node document initialization</title>
<meta name="author" title="Sergey G. Grekhov" href="mailto:sgrekhov@unipro.ru">
<meta name="reviewer" title="Aleksei Yu. Semenov" href="mailto:a.semenov@unipro.ru"> <!-- 2013-06-27 -->
<meta name="assert" content="Node document of the template content attribute must be a document">
<link rel="help" href="http://www.w3.org/TR/2013/WD-html-templates-20130214/#template-element">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src='../testcommon.js'></script>
<link rel="stylesheet" href="/resources/testharness.css">
</head>
<body>
<div id="log"></div>
<script type="text/javascript">

test(function() {
    var d = newHTMLDocument();
    var t = d.createElement('template');

    assert_equals(t.content.ownerDocument, d, 'Wrong template content owner document');

}, 'Node document of the template content attribute must be a document');

</script>
</body>
</html>
