<!DOCTYPE html>
<html>
<head>
<title>HTML Templates: BODY element cannot be content of TEMPLATE element</title>
<meta name="author" title="Sergey G. Grekhov" href="mailto:sgrekhov@unipro.ru">
<meta name="reviewer" title="" href="mailto:"> <!-- 2013-06-20 -->
<meta name="assert" content="BODY element cannot be content of TEMPLATE element">
<link rel="help" href="http://www.w3.org/TR/2013/WD-html-templates-20130214/#template-element">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src='../testcommon.js'></script>
<link rel="stylesheet" href="/resources/testharness.css">
</head>
<body>
<div id="log"></div>
<script type="text/javascript">

// Test empty BODY element
test(function() {
    var d = newHTMLDocument();
    var t = d.createElement('template');

    t.innerHTML = '<body></body>';

    d.body.appendChild(t);

    assert_equals(t.content.childNodes.length, 0, 'Template cannot contain BODY element');

}, 'BODY element cannot be content of TEMPLATE element (test empty BODY element)');

// Test BODY element only
test(function() {
    var d = newHTMLDocument();
    var t = d.createElement('template');

    t.innerHTML = '<body><div>Some content</div></body>';

    d.body.appendChild(t);

    assert_equals(t.content.childNodes.length, 1,
            'Wrong number of template content children');
    assert_equals(t.content.firstChild.nodeName, 'DIV',
            'Template should contain children of ignored BODY element');

}, 'BODY element cannot be content of TEMPLATE element (test BODY element only)');

// Test BODY element and some valid element after BODY tag
test(function() {
    var d = newHTMLDocument();
    var t = d.createElement('template');

    t.innerHTML = '<body><div <div id="dv1">Some content</div></body><div id="dv2">Some valid content</div>';

    d.body.appendChild(t);

    assert_equals(t.content.childNodes.length, 2, 
            'Wrong number of template content children');
    assert_true(t.content.querySelector('#dv1') != null,
            'Template should contain children of the ignored BODY element');
    assert_true(t.content.querySelector('#dv2') != null,
            'Template should contain valid element');

}, 'BODY element cannot be content of TEMPLATE element ' 
        + '(test BODY element and some valid element after BODY tag)');

// Test BODY element and some valid element before BODY tag
test(function() {
    var d = newHTMLDocument();
    var t = d.createElement('template');

    t.innerHTML = '<div id="dv1">Some valid content</div><body><div id="dv2">Some content</div></body>';

    d.body.appendChild(t);

    assert_equals(t.content.childNodes.length, 2,
            'Template cannot contain BODY element');
    assert_true(t.content.querySelector('#dv1') != null,
            'Template should contain valid element');
    assert_true(t.content.querySelector('#dv2') != null,
            'Template should contain children of the ignored BODY element');

}, 'BODY element cannot be content of TEMPLATE element '
        + '(test BODY element and some valid element before BODY tag)');

// Test nested template
test(function() {
    var d = newHTMLDocument();
    var t = d.createElement('template');

    t.innerHTML = '<template id="t2"><body><span>Body!<span></body></template>';

    d.body.appendChild(t);

    assert_equals(t.content.childNodes.length, 1,
            'Template should contain nested template');
    assert_true(t.content.querySelector('#t2') != null,
            'Template should contain nested element');

    var t2 = t.content.querySelector('#t2');

    assert_equals(t2.content.childNodes.length, 1,
            'Template cannot contain BODY element');
    assert_equals(t2.content.firstChild.nodeName, 'SPAN',
            'Template cannot contain BODY element');

}, 'BODY element cannot be content of TEMPLATE element (test nested template)');

</script>
</body>
</html>
