<!DOCTYPE html>
<html>
<head>
<title>HTML Templates: Template element as a descendant of the frameset element.</title>
<meta name="author" title="Aleksei Yu. Semenov" href="mailto:a.semenov@unipro.ru">
<meta name="reviewer" title="" href="mailto:"> <!-- 2013-06-20 -->
<meta name="assert" content="Template element can be a descendant of the frameset element">
<link rel="help" href="http://www.w3.org/TR/2013/WD-html-templates-20130214/#template-element">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src='../testcommon.js'></script>
<link rel="stylesheet" href="/resources/testharness.css">
</head>
<body>
<div id="log"></div>
<script type="text/javascript">

testInIFrame('../resources/template-descendant-frameset.html', function(ctx) {
    var d = ctx.iframes[0].contentDocument;

    var f = d.querySelector('frameset');

    assert_true(f.querySelector('template')!=null,
        'Template element should be a descendant of the frameset element');

}, 'Template element as a descendant of the frameset element (load from file)');

testInIFrame('../resources/template-descendant-frameset.html', function(ctx) {
    var d = ctx.iframes[0].contentDocument;

    var f = d.querySelector('frameset');

    f.innerHTML = '';
    assert_true(d.querySelector('template')==null,
        'Initial conditions are not satisfied');

    f.innerHTML = '<template>some text</template>';

    assert_true(f.querySelector('template')!=null,
        'Template element should be a descendant of the frameset element');

}, 'Template element as a descendant of the frameset element (using innerHTML)');

testInIFrame('../resources/template-descendant-frameset.html', function(ctx) {
    var d = ctx.iframes[0].contentDocument;

    var f = d.querySelector('frameset');

    var t = d.createElement('template');
    f.appendChild(t);

    assert_equals(f.querySelectorAll('template').length, 2,
        'Template element should be a descendant of the frameset element');

}, 'Template element as a descendant of the frameset element (using appendChild)');


</script>
</body>
</html>