<!DOCTYPE html>
<html>
<head>
<title>HTML Templates: HTML element cannot be content of TEMPLATE element</title>
<meta name="author" title="Sergey G. Grekhov" href="mailto:sgrekhov@unipro.ru">
<meta name="reviewer" title="" href="mailto:"> <!-- 2013-06-20 -->
<meta name="assert" content="HTML element cannot be content of TEMPLATE element">
<link rel="help" href="http://www.w3.org/TR/2013/WD-html-templates-20130214/#template-element">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src='../testcommon.js'></script>
<link rel="stylesheet" href="/resources/testharness.css">
</head>
<body>
<div id="log"></div>
<script type="text/javascript">

// HTML element only
test(function() {
    var d = newHTMLDocument();
    var t = d.createElement('template');

    t.innerHTML = '<html><body></body></html>';

    d.body.appendChild(t);

    assert_equals(t.content.childNodes.length, 0,
            'Template cannot contain HTML element');

}, 'HTML element cannot be content of TEMPLATE elemen (HTML element only)');

// Some valid element before HTML element
test(function() {
    var d = newHTMLDocument();
    var t = d.createElement('template');

    t.innerHTML = '<div id="dv">Some text</div><html><body></body></html>';

    d.body.appendChild(t);

    assert_equals(t.content.childNodes.length, 1,
            'Template cannot contain HTML element');
    assert_true(t.content.querySelector('#dv') != null,
            'Template should contain valid element');

}, 'HTML element cannot be content of TEMPLATE elemen '
    + '(some valid element before HTML element)');

// Some valid element after HTML element
test(function() {
    var d = newHTMLDocument();
    var t = d.createElement('template');

    t.innerHTML = '<html><body></body></html><div id="dv">Some text</div>';

    d.body.appendChild(t);

    assert_equals(t.content.childNodes.length, 1,
            'Template cannot contain HTML element');
    assert_true(t.content.querySelector('#dv') != null,
            'Template should contain valid element');

}, 'HTML element cannot be content of TEMPLATE elemen '
    + '(some valid element after HTML element)');

// Test nested template
test(function() {
    var d = newHTMLDocument();
    var t = d.createElement('template');

    t.innerHTML = '<template id="t2"><html><body></body></html></template>';

    d.body.appendChild(t);

    assert_equals(t.content.childNodes.length, 1,
            'Template should contain nested template');
    assert_true(t.content.querySelector('#t2') != null,
            'Template should contain nested element');

    var t2 = t.content.querySelector('#t2');

    assert_equals(t2.content.childNodes.length, 0,
            'Template cannot contain HTML element');

}, 'HTML element cannot be content of TEMPLATE elemen '
    + '(nested template)');

// Valid element inside HTML element
test(function() {
    var d = newHTMLDocument();
    var t = d.createElement('template');

    t.innerHTML = '<html><div id="dv">Some text</div></html>';

    d.body.appendChild(t);

    assert_equals(t.content.childNodes.length, 1,
            'Template cannot contain HTML element');
    assert_true(t.content.querySelector('#dv') != null,
            'Template should contain valid element');

}, 'HTML element cannot be content of TEMPLATE elemen '
    + '(valid element inside HTML element)');

// Valid element inside HTML and BODY element
test(function() {
    var d = newHTMLDocument();
    var t = d.createElement('template');

    t.innerHTML = '<html><body><div id="dv">Some text</div><body></html>';

    d.body.appendChild(t);

    assert_equals(t.content.childNodes.length, 1,
            'Template cannot contain HTML element');
    assert_true(t.content.querySelector('#dv') != null,
            'Template should contain valid element');

}, 'HTML element cannot be content of TEMPLATE elemen '
    + '(valid element inside HTML and BODY elements)');
    
//Valid element inside and between HTML and BODY element
test(function() {
    var d = newHTMLDocument();
    var t = d.createElement('template');

    t.innerHTML = '<html><span id="spn">Span</span><body><div id="dv">Some text</div><body></html>';

    d.body.appendChild(t);

    assert_equals(t.content.childNodes.length, 2,
            'Template cannot contain HTML element');
    assert_true(t.content.querySelector('#dv') != null,
            'Template should contain valid DIV element');

    assert_true(t.content.querySelector('#spn') != null,
            'Template should contain valid SPAN element');
}, 'HTML element cannot be content of TEMPLATE elemen '
    + '(valid element inside and between HTML and BODY elements)');
    
</script>
</body>
</html>
