<!DOCTYPE html>
<html>
<head>
<title>Custom element constructor sets value of is attribute to custom element type, if it is not equal to name</title>
<meta name="author" title="Sergey G. Grekhov" href="mailto:sgrekhov@unipro.ru">
<meta name="assert" content="If TYPE is not the same as NAME, set the value of ELEMENT's is attribute to TYPE">
<link rel="help" href="http://www.w3.org/TR/custom-elements/#instantiating-custom-elements">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src="../testcommon.js"></script>
<link rel="stylesheet" href="/resources/testharness.css">
</head>
<body>
<div id="log"></div>
<script type="text/javascript">
test(function() {
    var doc = newHTMLDocument();
    HTML5_ELEMENTS.forEach(function(tagName) {
        if (HTML5_DOCUMENT_ELEMENTS.indexOf(tagName) !== -1) {
            return;
        }
        var obj = doc.createElement(tagName);
        var name = 'x-c-' + tagName;
        var GeneratedConstructor = doc.registerElement(name, {
            prototype: newHTMLElementPrototype(name),
            extends: tagName});
        var customElement = new GeneratedConstructor();

        assert_equals(customElement.getAttribute('is'), name,
            'Value of the IS attribute should be set to type');
    });
}, 'Value of the IS attribute should be set to type if local name is not the same as type');


test(function() {
    var doc = newHTMLDocument();
    HTML5_ELEMENTS.forEach(function(tagName) {
        if (HTML5_DOCUMENT_ELEMENTS.indexOf(tagName) !== -1) {
            return;
        }
        var obj = doc.createElement(tagName);
        var name = 'x-d-' + tagName;
        var GeneratedConstructor = doc.registerElement(name, {
            prototype: newHTMLElementPrototype(name)});
        var customElement = new GeneratedConstructor();

        assert_equals(customElement.getAttribute('is'), null,
            'Value of the IS attribute should not be set');
    });
}, 'Value of the IS attribute should not be set to type if local name is the same as type');
</script>
</body>
</html>
