<!DOCTYPE html>
<html>
<head>
<title>HTML Templates: In body insertion mode: A start tag whose tag name is body</title>
<meta name="author" title="Sergey G. Grekhov" href="mailto:sgrekhov@unipro.ru">
<meta name="reviewer" title="" href="mailto:"> <!-- 2013-06-20 -->
<meta name="assert" content="If the stack of open elements has a template element in html scope then ignore <body> the token. (fragment or template contents case)">
<link rel="help" href="http://www.w3.org/TR/2013/WD-html-templates-20130214/#in-body-addition">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src='../../testcommon.js'></script>
<link rel="stylesheet" href="/resources/testharness.css">
</head>
<body>
<div id="log"></div>
<script type="text/javascript">

// <body> tag only
test(function () {
    var d = newHTMLDocument();
    
    d.body.innerHTML = '<template id="tmpl"><body></template>';
    
    var t = d.querySelector('#tmpl');
    
    assert_equals(t.content.childNodes.length, 0, 'Element must be ignored');
    
}, 'In body insertion mode: A start tag whose tag name is body (<body> tag only)');


// <body> tag with some text
test(function () {
    var d = newHTMLDocument();
    
    d.body.innerHTML = '<template id="tmpl"><body>Body text content</body></template>';
    
    var t = d.querySelector('#tmpl');
    
    assert_equals(t.content.querySelector('body'), null, '<body> element must be ignored');
    assert_equals(t.content.childNodes.length, 1, 'Text shouldn\'t be ignored');
    assert_equals(t.content.childNodes[0].nodeType, Node.TEXT_NODE, 'Text shouldn\'t be ignored');
    
}, 'In body insertion mode: A start tag whose tag name is body (<body> tag with some text)');


//<body> tag with some other elements
test(function () {
    var d = newHTMLDocument();
    
    d.body.innerHTML = '<template id="tmpl"><body>' 
        + '<div id="div1">DIV 1</div>' 
        + '<div id="div2">DIV 2</div>'
        + '</body></template>';
    
    var t = d.querySelector('#tmpl');
    
    assert_equals(t.content.querySelector('body'), null, '<body> element must be ignored');
    assert_equals(t.content.childNodes.length, 2, 'Only body tag should be ignored');
    assert_true(t.content.querySelector('#div1') != null, 'Children of <body tag shouldn\'t be ignored');
    assert_true(t.content.querySelector('#div2') != null, 'Children of <body tag shouldn\'t be ignored');
    
}, 'In body insertion mode: A start tag whose tag name is body ' 
        + '(<body> tag with some other elements)');


//nested template tag containing <body> tag with some other elements
test(function () {
    var d = newHTMLDocument();
    
    d.body.innerHTML = '<template id="tmpl1"><template id="tmpl2"><body>' 
        + '<div id="div1">DIV 1</div>' 
        + '<div id="div2">DIV 2</div>'
        + '</body></template></template>';
    
    var t = d.querySelector('#tmpl1').content.querySelector('#tmpl2');
    
    assert_equals(t.content.querySelector('body'), null, '<body> element must be ignored');
    assert_equals(t.content.childNodes.length, 2, 'Only body tag should be ignored');
    assert_true(t.content.querySelector('#div1') != null, 'Children of <body tag shouldn\'t be ignored');
    assert_true(t.content.querySelector('#div2') != null, 'Children of <body tag shouldn\'t be ignored');
    
}, 'In body insertion mode: A start tag whose tag name is body ' 
        + '(nested template tag containing <body> tag with some other elements');

</script>
</body>
</html>
