<!DOCTYPE html>
<html>
<head>
<title>If prototype interface inherits from HTMLElement then namespace must be HTML Namespace</title>
<meta name="author" title="Vasiliy Degtyarev" href="mailto:vasya@unipro.ru">
<meta name="assert" content="If prototype's interface inherits from HTMLElement, then NAMESPACE must be HTML Namespace">
<link rel="help" href="http://www.w3.org/TR/custom-elements/#registering-custom-elements">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src='../testcommon.js'></script>
<link rel="stylesheet" href="/resources/testharness.css">
</head>
<body>
<div id="log"></div>
<script type="text/javascript">

test(function() {
    var doc = newHTMLDocument();

    var xfoo = doc.createElementNS(HTML_NAMESPACE, 'x-foo');
    
    assert_equals(xfoo.namespaceURI, HTML_NAMESPACE, 'Wrong custom element namespace');
    
    var proto = Object.create(xfoo.constructor.prototype);
    var XFoo1 = doc.register('x-foo-1', {prototype: proto});
    
    var xfoo1 = new XFoo1();
    
    assert_equals(xfoo1.namespaceURI, HTML_NAMESPACE, 'Wrong namespace');
    
}, 'Namespace should be HTML for HTMLElement');

</script>
</body>
</html>
