<!DOCTYPE html>
<html>
<head>
<title>HTML Templates: Parsing XHTML: Node's node document</title>
<meta name="author" title="Sergey G. Grekhov" href="mailto:sgrekhov@unipro.ru">
<meta name="author" title="Aleksei Yu. Semenov" href="mailto:a.semenov@unipro.ru">
<meta name="assert" content="Parsing XHTML: Node's node document must be set to that of the element to which it will be appended">
<link rel="help" href="http://www.w3.org/TR/2013/WD-html-templates-20130214/#parsing-xhtml-documents">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src='../testcommon.js'></script>
<link rel="stylesheet" href="/resources/testharness.css">
</head>
<body>
<div id="log"></div>
<script type="text/javascript">



test(function() {
    var doc1 = newXHTMLDocument();
    var template = doc1.createElement('template');

    assert_equals(template.ownerDocument, doc1, 'Wrong template node owner document');
    assert_equals(template.content.ownerDocument, doc1,
            'Wrong template content owner document');

    var doc2 = newXHTMLDocument();

    doc2.body = doc2.createElement('body');
    doc2.body.appendChild(template);

    assert_equals(template.ownerDocument, doc2,
            'Template node owner document should be changed');

}, 'Parsing XHTML: Node\'s node document must be set to that of the element '
    + 'to which it will be appended. Test empty template');



test(function() {
    var doc1 = newXHTMLDocument();

    doc1.body = doc1.createElement('body');
    doc1.body.innerHTML = '<template id="tmpl"><div>Div content</div></template>';

    var template = doc1.querySelector('#tmpl');

    assert_equals(template.ownerDocument, doc1,
                    'Wrong template node owner document');

    var doc2 = newXHTMLDocument();
    var body2 = doc2.createElement('body');
    doc2.body = body2;
    doc2.body.appendChild(template);

    assert_equals(template.ownerDocument, doc2,
            'Template node owner document should be changed');

    assert_not_null(template.content,
            'Content attribute of TEMPLATE element shouldn\'t be null');

    var div = template.content.querySelector('div');
    assert_equals(template.content.ownerDocument, div.ownerDocument,
            'Wrong DIV node owner document');

}, 'Parsing XHTML: Node\'s node document must be set to that of the element '
    + 'to which it will be appended. Test not empty template');



test(function() {
    var doc1 = newXHTMLDocument();
    doc1.body = doc1.createElement('body');
    doc1.body.innerHTML = ''
            + '<template id="tmpl"><div>Div content</div> And some more text'
            + '<template id="tmpl2"><div>Template content</div></template>'
            + '</template>';

    var template = doc1.querySelector('#tmpl');

    assert_equals(template.ownerDocument, doc1, 'Wrong template node owner document');

    assert_not_null(template.content,
            'Content attribute of TEMPLATE element shouldn\'t be null');

    var nestedTemplate = template.content.querySelector('#tmpl2');

    assert_equals(nestedTemplate.ownerDocument, doc1,
            'Wrong nested template node owner document');

    var doc2 = newXHTMLDocument();
    var body2 = doc2.createElement('body');
    doc2.body = body2;
    doc2.body.appendChild(template);

    assert_equals(template.ownerDocument, doc2,
            'Template node owner document should be changed');

    assert_equals(nestedTemplate.ownerDocument, doc2,
            'Nested template node owner document should be changed');

    assert_not_null(nestedTemplate.content,
            'Content attribute of inner TEMPLATE element shouldn\'t be null');

    var div = nestedTemplate.content.querySelector('div');
    assert_equals(template.content.ownerDocument, div.ownerDocument,
            'Wrong DIV node owner document');

}, 'Parsing XHTML: Node\'s node document must be set to that of the element '
    + 'to which it will be appended. Test nested templates');



testInIFrame('../resources/template-child-nodes-div.xhtml', function(context) {
    var doc = context.iframes[0].contentDocument;

    var template = doc.querySelector('template');

    assert_equals(template.ownerDocument, doc, 'Wrong template node owner document');

    assert_not_null(template.content,
            'Content attribute of TEMPLATE element shouldn\'t be null');

    var div = template.content.querySelector('div');
    assert_equals(template.content.ownerDocument, div.ownerDocument,
            'Wrong DIV node owner document');

}, 'Parsing XHTML: Node\'s node document must be set to that of the element '
    + 'to which it will be appended. Test loading XHTML document from a file');



testInIFrame('../resources/template-child-nodes-nested.xhtml', function(context) {
    var doc = context.iframes[0].contentDocument;

    var template = doc.querySelector('template');

    assert_equals(template.ownerDocument, doc, 'Wrong template node owner document');

    var nestedTemplate = template.content.querySelector('template');

    assert_equals(nestedTemplate.ownerDocument, template.content.ownerDocument,
            'Wrong template node owner document');

    var div = nestedTemplate.content.querySelector('div');
    assert_equals(nestedTemplate.content.ownerDocument, div.ownerDocument,
            'Wrong DIV node owner document');

}, 'Parsing XHTML: Node\'s node document must be set to that of the element '
    + 'to which it will be appended. Test loading of XHTML document '
    + 'with nested templates from a file');


// load document from file and append template node to other XHTML document
testInIFrame('../resources/template-child-nodes-div.xhtml', function(context) {
    var doc1 = context.iframes[0].contentDocument;

    var template = doc1.querySelector('template');

    var doc2 = newXHTMLDocument();
    doc2.body = doc2.createElement('body');
    doc2.body.appendChild(template);

    assert_equals(template.ownerDocument, doc2, 'Wrong template node owner document');

    var div = template.content.querySelector('div');
    assert_equals(template.content.ownerDocument, div.ownerDocument,
            'Wrong DIV node owner document');

}, 'Parsing XHTML: Node\'s node document must be set to that of the element '
    + 'to which it will be appended. Test loading of XHTML document '
    + 'from a file and appending of template node to other XHTML document');



testInIFrame('../resources/template-child-nodes-nested.xhtml', function(context) {
    var doc1 = context.iframes[0].contentDocument;

    var template = doc1.querySelector('template');
    var nestedTemplate = template.content.querySelector('template');

    var doc2 = newXHTMLDocument();
    doc2.body = doc2.createElement('body');
    doc2.body.appendChild(template);

    assert_equals(template.ownerDocument, doc2, 'Wrong template node owner document');
    assert_equals(nestedTemplate.ownerDocument, template.content.ownerDocument,
            'Wrong template node owner document');


    var div = nestedTemplate.content.querySelector('div');
    assert_equals(nestedTemplate.content.ownerDocument, div.ownerDocument,
            'Wrong DIV node owner document');

}, 'Parsing XHTML: Node\'s node document '
    + 'Test loading of document with nested templates from server and '
    + 'append the template node to other XHTML document');

</script>
</body>
</html>
