<!DOCTYPE html>
<html>
<head>
<title>HTML Templates: Parsing XHTML: Node's node document</title>
<meta name="author" title="Sergey G. Grekhov" href="mailto:sgrekhov@unipro.ru">
<meta name="reviewer" title="" href="mailto:"> <!-- 2013-06-20 -->
<meta name="assert" content="Parsing XHTML: Node's node document must be set to the element to which it will be appended">
<link rel="help" href="http://www.w3.org/TR/2013/WD-html-templates-20130214/#parsing-xhtml-documents">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src='../testcommon.js'></script>
<link rel="stylesheet" href="/resources/testharness.css">
</head>
<body>
<div id="log"></div>
<script type="text/javascript">


// Empty template
test(function() {
    var d1 = newXHTMLDocument();
    var t = d1.createElement('template');

    assert_equals(t.ownerDocument, d1, 'Wrong template node owner document');
    assert_equals(t.content.ownerDocument, d1,
            'Wrong template content owner document');

    var d2 = newXHTMLDocument();

    d2.body = d2.createElement('body');
    d2.body.appendChild(t);

    assert_equals(t.ownerDocument, d2,
            'Template node owner document should be changed');

}, 'Parsing XHTML: Node\'s node document (empty template)');


// Not empty template
test(function() {
    var d1 = newXHTMLDocument();

    d1.body = d1.createElement('body');
    d1.body.innerHTML = '<template id="tmpl"><div>Div content</div></template>';

    var t = d1.querySelector('#tmpl');

    assert_equals(t.ownerDocument, d1,
                    'Wrong template node owner document');

    var d2 = newXHTMLDocument();
    var body2 = d2.createElement('body');
    d2.body = body2;
    d2.body.appendChild(t);

    assert_equals(t.ownerDocument, d2,
            'Template node owner document should be changed');

}, 'Parsing XHTML: Node\'s node document (not empty template)');


// Nested templates
test(function() {
    var d1 = newXHTMLDocument();
    d1.body = d1.createElement('body');
    d1.body.innerHTML = ''
            + '<template id="tmpl"><div>Div content</div> And some more text'
            + '<template id="tmpl2"><div>Template content</div></template>'
            + '</template>';

    var t = d1.querySelector('#tmpl');

    assert_equals(t.ownerDocument, d1, 'Wrong template node owner document');

    assert_true(t.content != null, 'Content attribute of TEMPLATE element shouldn\'t be null');
    var t2 = t.content.querySelector('#tmpl2');

    assert_equals(t2.ownerDocument, d1,
            'Wrong nested template node owner document');

    var d2 = newXHTMLDocument();
    var body2 = d2.createElement('body');
    d2.body = body2;
    d2.body.appendChild(t);

    assert_equals(t.ownerDocument, d2,
            'Template node owner document should be changed');

    assert_equals(t2.ownerDocument, d2,
            'Nested template node owner document should be changed');

}, 'Parsing XHTML: Node\'s node document (nested templates)');


// load XHTML document from file
testInIFrame('../resources/template-child-nodes-div.xhtml', function(ctx) {
    var d = ctx.iframes[0].contentDocument;

    var t = d.querySelector('template');

    assert_equals(t.ownerDocument, d, 'Wrong template node owner document');

}, 'Parsing XHTML: Node\'s node document (load XHTML document from file)');


// load document with nested templates from file
testInIFrame('../resources/template-child-nodes-nested.xhtml', function(ctx) {
    var d = ctx.iframes[0].contentDocument;

    var t1 = d.querySelector('template');

    assert_equals(t1.ownerDocument, d, 'Wrong template node owner document');

    var t2 = t1.content.querySelector('template');

    assert_equals(t2.ownerDocument, t1.content.ownerDocument, 'Wrong template node owner document');

}, 'Parsing XHTML: Node\'s node document (load document with nested templates from file)');


// load document from file and append template node to other XHTML document
testInIFrame('../resources/template-child-nodes-div.xhtml', function(ctx) {
    var d1 = ctx.iframes[0].contentDocument;

    var t = d1.querySelector('template');

    var d2 = newXHTMLDocument();
    d2.body = d2.createElement('body');
    d2.body.appendChild(t);

    assert_equals(t.ownerDocument, d2, 'Wrong template node owner document');

}, 'Parsing XHTML: Node\'s node document ' 
    + '(load document from file and append template node to other XHTML document)');


// load document with nested templates from server and
// append template node to other XHTML document
testInIFrame('../resources/template-child-nodes-nested.xhtml', function(ctx) {
    var d1 = ctx.iframes[0].contentDocument;

    var t1 = d1.querySelector('template');
    var t2 = t1.content.querySelector('template');

    var d2 = newXHTMLDocument();
    d2.body = d2.createElement('body');
    d2.body.appendChild(t1);

    assert_equals(t1.ownerDocument, d2, 'Wrong template node owner document');
    assert_equals(t2.ownerDocument, t1.content.ownerDocument, 
            'Wrong template node owner document');

}, 'Parsing XHTML: Node\'s node document ' 
    + '(load document with nested templates from server and '
    + 'append template node to other XHTML document)');

</script>
</body>
</html>
