<!DOCTYPE html>
<html>
<head>
<title>HTML Templates: ownerDocument property of the element in template</title>
<meta name="author" title="Sergey G. Grekhov" href="mailto:sgrekhov@unipro.ru">
<meta name="assert" content="ownerDocument property of the element in template is set to the template contents owner of the ownerDocument of the template element">
<link rel="help" href="http://www.w3.org/TR/2013/WD-html-templates-20130214/#creating-an-element-for-a-token">
<script src="/resources/testharness.js"></script>
<script src="/resources/testharnessreport.js"></script>
<script src='../../testcommon.js'></script>
<link rel="stylesheet" href="/resources/testharness.css">
</head>
<body>
<div id="log"></div>
<script type="text/javascript">


test(function () {
    var doc = newHTMLDocument();
    doc.body.innerHTML = '<div><template id="tmpl1"><div id="dv">DIV</div></template></div>';
    
    var template = doc.querySelector('#tmpl1');
        
    assert_equals(template.content.ownerDocument, template.ownerDocument, 
            'Wrong template content owner document');
    
    var div = template.content.querySelector('#dv');
    
    assert_equals(div.ownerDocument, template.ownerDocument, 
            'Wrong ownerDocument of the element in template');
    
}, 'Test ownerDocument property of the element in a template. '
    + 'Current DOCUMENT has no browsing context. Test template element inside the div');



test(function () {
    var doc = newHTMLDocument();
    doc.body.innerHTML = '<template id="tmpl1"><div id="dv">DIV</div></template>';
    
    var template = doc.querySelector('#tmpl1');
    
    assert_equals(template.content.ownerDocument, template.ownerDocument, 
            'Wrong template content owner document');
    
    var div = template.content.querySelector('#dv');
    
    assert_equals(div.ownerDocument, template.ownerDocument, 
            'Wrong ownerDocument of the element in template');
    
}, 'Test ownerDocument property of the element in a template. '
    + 'Current DOCUMENT has no browsing context. Test template element '
    + 'in the root of the body');



test(function () {
    var doc = newHTMLDocument();
    doc.head.innerHTML = '<template id="tmpl1"><div id="dv">DIV</div></template>';
    
    var template = doc.querySelector('#tmpl1');
    
    assert_equals(template.content.ownerDocument, template.ownerDocument, 
            'Wrong template content owner document');
    
    var div = template.content.querySelector('#dv');
    
    assert_equals(div.ownerDocument, template.ownerDocument, 
            'Wrong ownerDocument of the element in template');
    
}, 'Test ownerDocument property of the element in a template. '
    + 'Current DOCUMENT has no browsing context. Test template element '
    + 'in the root of the head');



test(function () {
    var doc = newHTMLDocument();
    
    doc.open();
    doc.write('<frameset><template id="tmpl1"><div id="dv">DIV</div></template></frameset>');
    doc.close();
    
    var template = doc.querySelector('#tmpl1');
    
    assert_equals(template.content.ownerDocument, template.ownerDocument, 
            'Wrong template content owner document');
    
    var div = template.content.querySelector('#dv');
    
    assert_equals(div.ownerDocument, template.ownerDocument, 
            'Wrong ownerDocument of the element in template');
    
}, 'Test ownerDocument property of the element in a template. '
    + 'Current DOCUMENT has no browsing context. Test template element '
    + 'in the root of the frameset');



test(function () {
    var doc = newHTMLDocument();
    doc.body.innerHTML = '<template id="tmpl1">' 
        + '<template id="tmpl2"><div id="dv">DIV</div></template></template>';
    
    var template = doc.querySelector('#tmpl1');
        
    assert_equals(template.content.ownerDocument, template.ownerDocument, 
            'Wrong template content owner document');
    
    var nestedTemplate = template.content.querySelector('#tmpl2');
    
    assert_equals(nestedTemplate.ownerDocument, template.content.ownerDocument, 
            'Wrong nested template owner document');
    
    assert_equals(nestedTemplate.content.ownerDocument, nestedTemplate.ownerDocument, 
            'Wrong nested template content owner document');
    
    var div = nestedTemplate.content.querySelector('#dv');
    
    assert_equals(div.ownerDocument, nestedTemplate.ownerDocument, 
            'Wrong ownerDocument');
    
}, 'ownerDocument property of the element in a template (nested template)');

</script>
</body>
</html>
